# Exercise Sheet 2
## Instructions
The instructions for this exercise sheet are the same as the previous, so this section is just a summary to serve as a reminder. Refer to the material on Moodle and at the start of Exercise Sheet 1 for more detail, and if you are unsure ask a tutor or post on the Q&A forum.

Complete each question in the code cell provided, which will usually contain some skeleton code to help you start the question. Unless specified otherwise you may change the skeleton code as you wish, but your code must pass the formatting tests in the following code cell – if the cell runs without errors then your code is eligible for submission. However these will only test the basics of the question. Your code will be subject to additional hidden tests which will determine your grade. Make sure you thoroughly test your own code to check you think it works correctly before submission.

**Note**: For any question which accepts a mutable input (e.g. a list), your code must not modify that input unless specified otherwise in the question.

## Questions
### Question 1
Write a function that takes a list of items and returns a new list with the order of items reversed, without using any inbuilt Python features that will do this for you (e.g. `.reverse()` or `reversed(…)`). As specified in the instructions, your function must not modify the contents of the original list.

<!-- grade_id: cell-3fe55db4c6d2fc90, points: 0, grade: false, solution: true -->
```python
def reversed_list(in_list):
    new = []
    for x in range(len(in_list)-1, -1, -1):
        new.append(in_list[x])
    return new
```

<!-- grade_id: cell-4788d108c8c7a005, points: 10, grade: true, solution: false -->
```python
assert(reversed_list([1, 2]) == [2, 1])
```

### Question 2
Now write a procedure that takes a list as an input and reverses its contents. So, the procedure should not return anything, and the inputted list *should* be modified. Of course, you still must not use inbuilt Python features which reverse lists.

<!-- grade_id: cell-26bdef0580e1c19c, points: 0, grade: false, solution: true -->
```python
def reverse_list(in_list):
    # YOUR CODE HERE
    new = []
    for x in range(len(in_list)-1, -1, -1):
        new.append(in_list[x])
    for x in range(len(in_list)):
        in_list[x] = new[x]
    
```

<!-- grade_id: cell-fde78ec384676dc5, points: 10, grade: true, solution: false -->
```python
test_in = [1, 2]
reverse_list(test_in)
assert(test_in == [2, 1])
```

### Question 3
Given two lists `[a1, a2, a3, …]` and `[b1, b2, b3, …]` of the same length, return a list containing pairs of items as tuples by position: a tuple containing the first item from each list, a tuple containing the second item, and so on. In other words, return `[(a1, b1), (a2, b2), …]`.

This is quite a useful feature when you want to iterate over two collections at once, and it also has a inbuilt function called `zip`! Of course, you cannot use `zip` (or other similar functions) in your submission.

<!-- grade_id: cell-2b3c95462f105301, points: 0, grade: false, solution: true -->
```python
def merge(list1, list2):
    zips = []
    for x in range(0, len(list1)):
        zips.append((list1[x], list2[x]))
    return zips
```

<!-- grade_id: cell-328705de8fb298ac, points: 10, grade: true, solution: false -->
```python
assert(merge([1, 2], [3, 4]) == [(1, 3), (2, 4)])
```

### Question 4
Write another function like the one from the previous question. It should take two *non-empty* lists and create a list of tuples. The two input lists do not need to be the same length. If they are not the same length, *wrap around* the shorter list, i.e. start from the beginning again. (Hint: don't forget the *modulo* operator.) As usual, do not use `zip` or other similar inbuilt functions which implement this functionality already.

So for example, given the inputs `[1, 2, 3, 4, 5]`, and `['a', 'b']`, the result should be: <br>
`[(1, 'a'), (2, 'b'), (3, 'a'), (4, 'b'), (5, 'a')]`.

<!-- grade_id: cell-4e38ae067a6ac9c4, points: 0, grade: false, solution: true -->
```python
def merge_wrap(list1, list2):
    zips = []
    longer_list = list1
    if len(longer_list) < len(list2):
        longer_list = list2
    for x in range(0, len(longer_list)):
        zips.append((list1[x%len(list1)], list2[x%len(list2)]))
    return zips
```

<!-- grade_id: cell-b3b03350120ec08f, points: 10, grade: true, solution: false -->
```python
assert(merge_wrap([1, 2], [3, 4, 5]) == [(1, 3), (2, 4), (1, 5)])
```

### Question 5
Now one more variation: this time you must write a version of the merge and wrap function which takes *any number* of nonempty lists as input, using the *varargs* syntax you saw in the material about tuples. The return value should be a single list containing a number of tuples equal to the length of the longest list from all of the arguments, where each tuple has a number of elements equal to the number of arguments. Hopefully this is obvious, but you must not use `zip` or any other similar inbuilt functions which do the work for you.

See an example in the test cell below.

<!-- grade_id: cell-855bf4e1dfa15c7b, points: 0, grade: false, solution: true -->
```python
def merge_wrap_n(*args):
    # YOUR CODE HERE
    zips = []
    longest_list = args[0]
    for x in args:
        if len(x) > len(longest_list):
            longest_list = x
    for x in range(0, len(longest_list)):
        new_pair = []
        for i in args:
            new_pair.append(i[x%len(i)])
        zips.append(tuple(new_pair))
    return zips
```

<!-- grade_id: cell-bb2620ea2d8898af, points: 10, grade: true, solution: false -->
```python
assert(merge_wrap_n([1, 2], [3], [4, 5, 6]) == [(1, 3, 4), (2, 3, 5), (1, 3, 6)])
```

### Question 6
Write a function which takes a list of numbers of length $\ge 2$, and returns a tuple containing the *arithmetic mean* and the *unbiased sample variance* of those numbers.

i.e. given a list of items $[x_1, x_2, \dots, x_N]$, return a tuple $(m, s^2)$ where:

\begin{align} 
m &= \frac{1}{N} \sum_{i=1}^{N} x_i \\
s^2 &= \frac{1}{N-1} \sum_{i=1}^{N} (x_i - m)^2 .
\end{align}

The symbol $\sum$ means to add up all the values in this range for the variable specified. In other words, take the value where $i=1$, add the value where $i=2$, and so on until $i=N$. So the arithmetic mean can be explained in words as: add up all the values in the list, then divide by the total number of values ($N$). For the variance you need to subtract the mean from each value and square the result, add this up, then divide by $N-1$. You are welcome to use mathematically equivalent reformulations if you prefer (and it may be more efficient to do so), just ensure you match the outputs of these formulas exactly.

You must do the calculations manually, do not use inbuilt functions like `sum`, any `math` functions, etc.

<!-- grade_id: cell-bece9af02711a741, points: 0, grade: false, solution: true -->
```python
def mean_variance(numbers):
    m = 0 
    s = 0 
    
    for x in numbers:
        m += x
    m /= len(numbers)
    for x in numbers:
        s += (x-m)** 2
    s /= (len(numbers)-1)
    return (m, s)
```

<!-- grade_id: cell-8b8d9bd745e70c03, points: 10, grade: true, solution: false -->
```python
assert(mean_variance([1, 2, 3]) == (2.0, 1.0))
```

### Question 7
Write a function which takes a string as an input and returns a dictionary which maps from each letter a-z (the keys), *ignoring case*, to the number of times this letter appears in the string (the values), *if the letter appears in the string* (all counts should be $\ge 1$). The keys to the dictionary must all be *lower case*. Ignore all punctuation or whitespace.

<!-- grade_id: cell-f38552901de3cd4a, points: 0, grade: false, solution: true -->
```python
def frequency_analysis(string):
    # YOUR CODE HERE
    freq = {}
    for x in string:
        x = x.lower()
        if (x < "a") or (x > "z"):
            continue
        if x in freq.keys():
            freq[x] += 1
        else:
            freq[x] = 1
    return freq
```

<!-- grade_id: cell-3f5ae3bb369a1544, points: 10, grade: true, solution: false -->
```python
assert(frequency_analysis("Hello world!") == {'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1})
```

### Question 8
Write a function which takes a string as an input and returns a dictionary where the keys are the *words* from the string and the values are the *counts*, how often each word appears in the string. A word is a sequence of consecutive letters which is broken by a space, the end of the string, or any of the following punctuation symbols: `.,?!` (no other punctuation will be included in the input string). Ignore case when counting words: the dictionary's keys must be entirely in lower case. Ignore consecutive spaces or punctuation.

You are *encouraged* to make use of the various [string methods](https://docs.python.org/3/library/stdtypes.html#string-methods) available in Python.

<!-- grade_id: cell-441b97aa02b39592, points: 0, grade: false, solution: true -->
```python
def count_words(string):
    current_word = ""
    words = {}
    for x in string:
        x = x.lower()
        if (x < "a") or (x > "z"):
            if current_word == "":
                continue
            else:
                if current_word in words.keys():
                    words[current_word] += 1
                else:
                    words[current_word] = 1
            current_word = ""
            continue
        current_word += x

    if current_word != "":
        if current_word in words.keys():
            words[current_word] += 1
        else:
            words[current_word] = 1

    return words
```

Output:
```
{'hello': 1, 'subeen': 2, 'a': 1, 'hell': 1}
```

<!-- grade_id: cell-1d02a1dcee1ca237, points: 10, grade: true, solution: false -->
```python
assert(count_words("Hello world!") == {"hello": 1, "world": 1})
assert(count_words("The Bart, the") == {"the": 2, "bart": 1})
```

### Question 9
In this question you will implement the [sieve of Eratosthenes](https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes), a method for finding prime numbers. A prime number is any integer which is divisible by exactly two numbers: $1$ and itself. The basic idea is that you start with a list (or equivalent) of all the numbers from $2$ to $N$, then you repeat the following two steps:
* remove the first element from the list, this is your next prime number, call it $p$
* remove every element from the rest of the which is a multiple of $p$

By repeating these steps you fill find every prime number up to $N$. There are various optimisations you can perform, but this basic idea is sufficient for this exercise.

For this question, write a function which takes a single integer `n` and returns a *set* containing all of the prime numbers up to *and including* `n`.

<!-- grade_id: cell-657584a91cd8771d, points: 0, grade: false, solution: true -->
```python
def sieve(n):
    # YOUR CODE HERE

    primes = set()
    numbers = list(range(2, n+1))
    finished = False
    if n < 2:
        return primes
    while not finished:
        p = numbers[0]
        primes.add(p)
        for x in numbers:
            if x % p == 0:
                numbers.remove(x)
        if len(numbers) == 0:
            finished = True

    return primes
```

<!-- grade_id: cell-8203fc04b66a67e8, points: 10, grade: true, solution: false -->
```python
assert(sieve(10) == {2, 3, 5, 7})
```

### Question 10
In this question you need to write three functions, all of which return 2D lists with a given pattern made up of single-character strings containing either a hash symbol `'#'` or a space `' '`.

The three patterns are:
* `rectangle(n, m)`, where $n\ge2$ and $m\ge2$, must produce an `n` by `m` rectangle outline, e.g:
  ```
  rectangle(3, 3) returns
      [['#', '#', '#'], 
       ['#', ' ', '#'], 
       ['#', '#', '#']]
       
  rectangle(4, 5) returns
      [['#', '#', '#', '#', '#'],
       ['#', ' ', ' ', ' ', '#'],
       ['#', ' ', ' ', ' ', '#'],
       ['#', '#', '#', '#', '#']]
  ```
* `diamond(n)`, where $n\ge3$, must produce an `n` by `n` filled diamond, where odd sized diamonds are one hash wide at their points, and even shaped diamonds are two hashes wide, e.g.:
  ```
  diamond(5) returns
      [[' ', ' ', '#', ' ', ' '],
       [' ', '#', '#', '#', ' '],
       ['#', '#', '#', '#', '#'],
       [' ', '#', '#', '#', ' '],
       [' ', ' ', '#', ' ', ' ']]
       
  diamond(6) returns
      [[' ', ' ', '#', '#', ' ', ' '],
       [' ', '#', '#', '#', '#', ' '],
       ['#', '#', '#', '#', '#', '#'],
       ['#', '#', '#', '#', '#', '#'],
       [' ', '#', '#', '#', '#', ' '],
       [' ', ' ', '#', '#', ' ', ' ']]
  ```
* `chequerboard(n, m, c_h, c_w)`, all arguments $\ge1$, returns an `n` by `m` chequerboard where each chequer (region) is of size `c_h` by `c_w`, the top left cell is *always* filled, and *only* the right hand and bottom edge may contain incomplete chequers, e.g.:
  ```
  chequerboard(6, 8, 2, 2) returns
      [['#', '#', ' ', ' ', '#', '#', ' ', ' '],
       ['#', '#', ' ', ' ', '#', '#', ' ', ' '],
       [' ', ' ', '#', '#', ' ', ' ', '#', '#'],
       [' ', ' ', '#', '#', ' ', ' ', '#', '#'],
       ['#', '#', ' ', ' ', '#', '#', ' ', ' '],
       ['#', '#', ' ', ' ', '#', '#', ' ', ' ']]
      
  chequerboard(5, 7, 2, 3) returns
      [['#', '#', '#', ' ', ' ', ' ', '#'],
       ['#', '#', '#', ' ', ' ', ' ', '#'],
       [' ', ' ', ' ', '#', '#', '#', ' '],
       [' ', ' ', ' ', '#', '#', '#', ' '],
       ['#', '#', '#', ' ', ' ', ' ', '#']]
  ```
  
  
For this question, `rectangle(n, m)` is worth a maximum of 2 marks, `diamond(n)` is worth 3 marks, and `chequerboard(…)` is worth 5 marks.

<!-- grade_id: cell-1d1b2cff91c4b309, points: 0, grade: false, solution: true -->
```python
def rectangle(n, m):
    frame = []
    for r in range(0, n):
        currentrow = []
        for c in range(0, m):
            if (r==0) or (r==n-1):
                currentrow.append('#')
            else:
                if (c==0) or (c==m-1):
                    currentrow.append('#')
                else:
                    currentrow.append(' ')   
        frame.append(currentrow)
    return frame

def diamond(n):
    # YOUR CODE HERE
    frame = []
    for r in range(0, n):
        current_row = []
        for c in range(0, n):
            current_row.append('#')
        frame.append(current_row)
    
    corners = (n-1)//2
    t = 1
    goneOnce = False
    for r in range(0, n):
        if goneOnce:
            goneOnce = False
            continue
        for x in range(0, len(frame[r])):
            if x > corners-1:
                break
            else:
                frame[r][x] = ' '
                frame[r][-1-x] = ' '

        corners -= t
        if corners == 0:
            if n%2 == 0:
                goneOnce = True
            t=-1
        
    return frame 
    
def chequerboard(n, m, c_h, c_w):
    # YOUR CODE HERE
    frame = []
    for r in range(0, n):
        current_row = []
        for c in range(0, m):
            if (((c // c_w) % 2 == 0) and ((r//c_h) % 2 == 0)) or (((r//c_h) % 2 == 1) and ((c//c_w) % 2 == 1)):
                current_row.append('#')
            else:
                current_row.append(' ')
        frame.append(current_row)

    return frame
```

<!-- grade_id: cell-699d6450bbd8bbd9, points: 10, grade: true, solution: false -->
```python
rect = [['#', '#', '#'], 
        ['#', ' ', '#'], 
        ['#', '#', '#']]

diam = [[' ', '#', ' '], 
        ['#', '#', '#'], 
        [' ', '#', ' ']]

cheq = [['#', ' ', '#'], 
        [' ', '#', ' '], 
        ['#', ' ', '#']]

assert(rectangle(3, 3) == rect)
assert(diamond(3) == diam)
assert(chequerboard(3, 3, 1, 1) == cheq)
```

```python

```